# Autoencoders

```python
import torch
import torch.utils.data
from torch import nn, optim
from torch.nn import functional as F
from torchvision import datasets, transforms
from torchvision.utils import save_image
from torchvision.utils import make_grid
from torch.utils.tensorboard import SummaryWriter


seed=1
device = torch.device("cuda" if torch.cuda.is_available() else "cpu")
```

# Dataset
MNIST dataset will be used

```python
from IPython.core.display import HTML
from markdown import markdown
import os

print('Run the following command in terminal:')
print('tensorboard --logdir runs --bind_all')

markdown_string = f"Tensorboard URL: [https://mlhep.coresearch.club/{os.environ['COCALC_PROJECT_ID']}/server/6006/](https://mlhep.coresearch.club/{os.environ['COCALC_PROJECT_ID']}/server/6006/)"
writer = SummaryWriter('runs/')

HTML("<p>{}</p>".format(markdown(markdown_string)))
```

```python
batch_size=128
kwargs = {'num_workers': 1, 'pin_memory': True} if torch.cuda.is_available() else {}

mnist_transforms = transforms.Compose([ 
    transforms.ToTensor(), # PIL Image -> Tensor
#     transforms.Lambda(lambda x: x/255.),
])

train_loader = torch.utils.data.DataLoader(
    datasets.MNIST('../../../share', train=True,
                   transform=mnist_transforms),
    batch_size=batch_size, shuffle=True, **kwargs)

test_loader = torch.utils.data.DataLoader(
    datasets.MNIST('../../../share', train=False, transform=mnist_transforms),
    batch_size=batch_size, shuffle=True, **kwargs)
```

```python
image[indx]
```

```python
image, label = next(iter(train_loader))
image.shape
%matplotlib inline
import matplotlib.pyplot as plt
indx = 0
plt.title(f'MNIST: {label[indx].item()}')
plt.imshow(image[indx].squeeze(0))
```

## Autoencoder model
![](http://bjlkeng.github.io/images/autoencoder_structure.png)

<!-- grade_id: 6404ff, points: 0, grade: false, solution: true -->
```python
class AE(nn.Module):
    def __init__(self, hidden_size=20):
        super(AE, self).__init__()
        self.hidden_size = hidden_size
        
        # Define self.encoder and self.decoder ()
        
        # your code here
        raise NotImplementedError

    def encode(self, x):
        # your code here
        raise NotImplementedError
    def decode(self, z): return self.decoder(z).view(-1,28,28)
    def forward(self, x):
        # your code here
        raise NotImplementedError
    def sample(self, size):
        # your code here
        raise NotImplementedError
    @property
    def device(self): return next(self.parameters()).device


model = AE().to(device)
optimizer = optim.Adam(model.parameters(), lr=1e-3)
```

<!-- grade_id: eba360, points: 1, grade: true, solution: false -->
```python
for batch_idx, (data, _) in enumerate(train_loader):
    data_recon = model(data.to(device))
    assert data_recon.shape == torch.Size([128,28,28])
    break
```

## Train model

<!-- grade_id: 0e68e0, points: 0, grade: false, solution: true -->
```python
def recon_loss(recon_x, x):
    # your code here
    raise NotImplementedError
    return MSE


log_interval=10
epochs=30

def train(epoch):
    model.train()
    train_loss = 0
    # your code here
    raise NotImplementedError
    train_loss /= len(train_loader.dataset)
    writer.add_scalar('ae/train_loss', train_loss, epoch)

    print('====> Epoch: {} Average loss: {:.4f}'.format(epoch, train_loss))


def test(epoch):
    model.eval()
    test_loss = 0
    # your code here
    raise NotImplementedError
    test_loss /= len(test_loader.dataset)
    writer.add_scalar('ae/test_loss', test_loss)
    return test_loss

for epoch in range(1, epochs + 1):
    train(epoch)
    test_loss = test(epoch)
    with torch.no_grad():
        sample_size=64
        sample = model.sample(sample_size).cpu()
        img = make_grid(sample.view(-1,1,28,28))
        writer.add_image('ae/test_image', img, epoch)
```

<!-- grade_id: a31e56, points: 2, grade: true, solution: false -->
```python

```

```python
img = (img.permute((1,2,0))*255).detach().cpu().numpy().astype('uint8').squeeze()
plt.imshow(img)
```

# Denoising autoencoder

<!-- grade_id: e74702, points: 0, grade: false, solution: true -->
```python
model = AE().to(device)
optimizer = optim.Adam(model.parameters(), lr=1e-3)

def recon_loss(recon_x, x):
    # your code here
    raise NotImplementedError
    return MSE


# log_interval=10
epochs=30

def train(epoch,var=1e-2):
    model.train()
    train_loss = 0
    # Add noise to image in training stage with "var" variance
    # your code here
    raise NotImplementedError
    
    train_loss /= len(train_loader.dataset)
    writer.add_scalar('dae/train_loss', train_loss)

    print('====> Epoch: {} Average loss: {:.4f}'.format(epoch, train_loss))



def test(epoch):
    model.eval()
    test_loss = 0
    # your code here
    raise NotImplementedError
    test_loss /= len(test_loader.dataset)
    writer.add_scalar('dae/test_loss', test_loss)
    return test_loss

for epoch in range(1, epochs + 1):
    train(epoch)
    test_loss = test(epoch)
    with torch.no_grad():
        sample_size=64
        sample = model.sample(sample_size).cpu()
        img = make_grid(sample.view(-1,1,28,28))
        writer.add_image('dae/test_image', img, epoch)
        
```

<!-- grade_id: 5bf284, points: 2, grade: true, solution: false -->
```python

```

```python
img = (img.permute((1,2,0))*255).detach().cpu().numpy().astype('uint8').squeeze()
plt.imshow(img)
```

# Sparse autoencoder

<!-- grade_id: bf98a3, points: 0, grade: false, solution: true -->
```python
# Implement Sparse autoencoder with L1 regularization on each intermediate activation

class SparseAE(nn.Module):
    def __init__(self, hidden_size=20):
        super(SparseAE, self).__init__()
        self.hidden_size = hidden_size
        
        # Define self.encoder and self.decoder ()
        
        # your code here
        raise NotImplementedError

    def encode(self, x):
        # your code here
        raise NotImplementedError
    def decode(self, z): return self.decoder(z).view(-1,28,28)
    def forward(self, x):
        # your code here
        raise NotImplementedError
    def sample(self, size):
        # your code here
        raise NotImplementedError
    @property
    def device(self): return next(self.parameters()).device

model = SparseAE().to(device)
optimizer = optim.Adam(model.parameters(), lr=1e-3)

def recon_loss(recon_x, x):
    # your code here
    raise NotImplementedError
    return MSE


# log_interval=10
epochs=30
var = 1e-2

def train(epoch):
    model.train()
    train_loss, train_loss_reg = 0, 0
    # your code here
    raise NotImplementedError
        
    train_loss /= len(train_loader.dataset)
    train_loss_reg /= len(train_loader.dataset)
    writer.add_scalar('dae/train_loss', train_loss)
    writer.add_scalar('dae/train_loss_reg', train_loss_reg)
            
    print('====> Epoch: {} Average loss: {:.4f} Reg loss: {:.4f}'.format(epoch, train_loss, train_loss_reg))


def test(epoch):
    model.eval()
    test_loss, test_loss_reg = 0, 0
    # your code here
    raise NotImplementedError
    test_loss /= len(test_loader.dataset)
    test_loss_reg /= len(test_loader.dataset)
    writer.add_scalar('dae/test_loss', test_loss)
    return test_loss, test_loss_reg

for epoch in range(1, epochs + 1):
    train(epoch)
    test_loss, test_loss_reg = test(epoch)
    with torch.no_grad():
        sample_size=64
        sample = model.sample(sample_size).cpu()
        img = make_grid(sample.view(-1,1,28,28))
        writer.add_image('sae/test_image', img, epoch)
```

<!-- grade_id: edd62b, points: 3, grade: true, solution: false -->
```python

```

```python
img = (img.permute((1,2,0))*255).detach().cpu().numpy().astype('uint8').squeeze()
plt.imshow(img)
```

# Links

[Simple explanation of AE](http://bjlkeng.github.io/posts/autoregressive-autoencoders/)

[Sparse autoencoder](https://debuggercafe.com/sparse-autoencoders-using-l1-regularization-with-pytorch/)

[Denoising autoencoder](https://towardsdatascience.com/denoising-autoencoders-explained-dbb82467fc2)

```python

```